Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# DSA 2018 Nyeri Preparatory Notebook
By Ciira Maina


In preparation for DSA 2018 Nyeri, we would like potential participants to complete a number of exercises in probability, machine learning and programming to ensure that they have the necessary prerequisite knowledge to attend the summer school. You will be required to submit notebooks with solutions to these exercises during the application process.

In this first exercise we will require you to download a dataset and perform computations on the data. These data are from a paper in 1966 by Cox and Lewis and report the time difference between nerve pulses on a nerve fibre. 799 observations are reported. These data are used for some examples in the text ["All of Statistics"](http://www.stat.cmu.edu/~larry/all-of-statistics/) by Larry Wasserman.

The data are available [here](http://www.stat.cmu.edu/~larry/all-of-statistics/=data/nerve.dat)

## Obtain the data

Write code to obtain the data from the website above and store it in a one dimensional array of floating point numbers.

<!-- grade_id: cell-8b789c350a9f25d6, points: 5, grade: true, solution: true -->
```python
import urllib.request
import numpy as np


# Your code here
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-88e09d936386e957, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN_TESTS
assert len(nerve_data) == 799
### END HIDDEN_TESTS
```

## Preliminary Visualisation

Plot a histogram of the data. Ensure you label your axes.

<!-- grade_id: cell-eb77be4ffc0854f8, points: 5, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt
%matplotlib inline

#Your code here
# YOUR CODE HERE
raise NotImplementedError()
```

## Preliminary analysis

The cumulative distribution function of a random variable $\mathbf{X}$ is given by

$
\begin{equation}
F_X(x)=P(\mathbf{X}\leq x)
\end{equation}$

If we obtain $n$ observations $X_1,\ldots,X_n$ from this distribution, the empirical distibution function is given by

$
\begin{equation}
\hat{F}_n(x)=\frac{\sum_{i=1}^n\mathbf{I}(X_i\leq x)}{n}
\end{equation}$

where
$
 \begin{equation}
       \mathbf{I}(X_i\leq x) = 
        \begin{cases}
            1 & \text{if $X_i\leq x $} \\
            0 & \text{if $X_i> x$} 
        \end{cases}
    \end{equation}
$
* Plot the empirical distribution function of the nerve pulse data
* Estimate the probability that the wait time between nerve pulses is less than $0.3$ - We will call this P1
* Estimate the probability that the wait time between nerve pulses is between $0.1$ and $0.3$ - We will call this P2



<!-- grade_id: cell-ce56deb8dc558f32, points: 5, grade: true, solution: true -->
```python
x=np.linspace(0,np.max(nerve_data),100)
cdf=np.zeros(100)

# Your code here
# YOUR CODE HERE
raise NotImplementedError()

#Your code here
#P1=
#P2=
```

<!-- grade_id: cell-96dbfa4daec9607c, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN_TESTS
assert abs(P1-0.760951188986) < 1e-6
assert abs(P2-0.377972465582) < 1e-6
### END HIDDEN_TESTS
```

## Estimating properties of the distribution

We can estimate properties of the true distribution of the data $F_X(x)$ using the empirical distribution function $\hat{F}_n(x)$. To do this we can use "plug in" estimators. Here we will estimate the mean, variance and skewness. The expressions for the "plug in" estimators for these quantities are
* Mean: $\hat{\mu}=\frac{1}{n}\sum_{i=1}^nX_i$
* Variance: $\hat{\sigma}^2=\frac{1}{n}\sum_{i=1}^n(X_i-\hat{\mu})^2$
* Skewness: $\hat{\kappa}=\frac{\frac{1}{n}\sum_{i=1}^n(X_i-\hat{\mu})^3}{\hat{\sigma}^3}$

Compute the plug in estimators of the mean, variance and skewness for the nerve pulse wait time data.

<!-- grade_id: cell-a11ec7e089beddcb, points: 5, grade: true, solution: true -->
```python
# Your code here
#mu_hat=
#var_hat=
#kappa_hat=
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```